# Sistemas mecánicos

El objetivo de esta práctica es analizar sistemas mecánicos a traves de las herramientas matemáticas de la ingeniería de control.

Empecemos modelando el siguiente sistema mecánico:

![](./imagenes/mra.png)

Si dibujamos nuestro diagrama de cuerpo libre y escribimos la sumatoria de fuerzas en $x$, obtendremos:

$$
\sum F_x = F - F_R - F_A = ma
$$

en donde $F$ es la fuerza aplicada hacia la derecha, $F_R$ es la fuerza de reacción en el resorte y $F_A$ es la fuerza de reacción en el amortiguador.

Si ahora tomamos en cuenta que:

$$
\begin{align}
F_R &= k x \\
F_A &= c v = c\dot{x} \\
ma  &= m \ddot{x}
\end{align}
$$

podemos escribir esta sumatoria de fuerzas como:

$$
F - kx - c\dot{x} = m \ddot{x}
$$

y sacando la transformada de Laplace, y factorizando terminos comúnes:

$$
F(s) = X(s)\left[ ms^2 + cs + k \right]
$$

Por lo que, cuando consideramos a $F(s)$ como la entrada de nuestro sistema y a $X(s)$ como la salida de nuestro sistema, podemos obtener la función de transferencia:

$$
\frac{X(s)}{F(s)} = \frac{1}{ms^2 + cs + k}
$$

y simular el comportamiento de este:

```python
from control import tf, step_response, root_locus
from numpy import linspace
from matplotlib.pyplot import plot
```

```python
m = 1200/4
c = 1500
k = 15000
G = tf([1], [m, c, k])
```

```python
ts = linspace(0, 10, 500)
t, y = step_response(G, ts)
```

```python
plot(t, y);
```

Sin embargo, tenemos que revisar una cosa. Los datos ingresados a este sistema se obtuvieron de valores comerciales para la suspensión de un automovil tipo sedán; sin embargo la función ```step_response``` simula el comportamiento del sistema para una entrada unitaria (en este caso $1N$), por lo que para que esta simulación tenga relevancia tenemos que amplificar esta entrada.

Se propone una entrada de $1100N$ que corresponse al peso de un hombre pesado, lo que esperaremos entonces es un movimiento como el que pasa cuando un hombre pesado se sube a un sedán:

```python
ts = linspace(0, 10, 500)
t, y = step_response(1100*G, ts)
```

```python
plot(t, y);
```

Y ahora obtenemos una simulación con el comportamiento que esperariamos de la suspensión de un coche cuando un hombre pesado se sube en el; la suspensión deja de moverse despues de unos $3s$ y se estabiliza en un valor de aproximadamente $0.07m$, es decir $7cm$ despues de haberse comprimido casi $10cm$  y regresar unas $3$ o $4$ veces.

---
## Ejercicio

* Define un sistema ```G1``` con un amortiguador con constante $c=0\frac{Ns}{m}$, un resorte con constante $k=100\frac{N}{m}$ y una masa de $10kg$.

<!-- grade_id: cell-7ce6a17c7ef1d899, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-cf63638779f82747, points: 1.5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

assert_allclose(G1.dcgain(), [0.01], 2)
assert_allclose(G1.pole(), [0.+3.16j,  0.-3.16j], 2)
assert_allclose(G1.zero(), [], 2)
```

```python
G1.zero()
```

* Simula el comportamiento de este sistema para una fuerza aplicada de $5N$ del tiempo $0s$ al tiempo $15s$.

<!-- grade_id: cell-2b292a828dad999c, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
plot(t, y);
```

<!-- grade_id: cell-e3ee246bb62b0395, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_equal

assert_equal(ts[0], 0)
assert_equal(ts[-1], 15)
assert_almost_equal(max(y), 0.02, 4)
assert_almost_equal(min(y), 0.0, 4)
```

* Define un sistema ```G2``` con un amortiguador con constante $c=10\frac{Ns}{m}$, un resorte con constante $k=0\frac{N}{m}$ y una masa de $10kg$.

<!-- grade_id: cell-7e1f420bb5263764, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-23657621642284d7, points: 1.5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
from numpy import inf

assert_allclose(G2.dcgain(), [inf])
assert_allclose(G2.pole(), [-1, 0])
assert_allclose(G2.zero(), [], 2)
```

* Simula el comportamiento de este sistema para una fuerza aplicada de $5N$ del tiempo $0s$ al tiempo $20s$.

<!-- grade_id: cell-49b40d70c45492c9, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
plot(t, y);
```

<!-- grade_id: cell-646252622d3a3923, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_almost_equal, assert_equal

assert_equal(ts[0], 0)
assert_equal(ts[-1], 20)
assert_almost_equal(max(y), 1.9, 4)
assert_almost_equal(min(y), 0.0, 4)
```

---

Una vez que hemos comprobado la manera de simular estos sistemas mecánicos, podemos pasar al siguiente paso, predecir su comportamiento a partir de la pura función de transferencia.

La función de transferencia tiene varias caracteristicas de las que no hemos hablado, por ejemplo los polos del sistema:

```python
G1 = tf([1], [1, 1, 1])
G1.pole()
```

Estos polos son obtenidos al resolver la ecuación formada al igualar el denominador de la función de transferencia a $0$.

> Este denominador es llamado **Polinomio caracteristico del sistema** ya que es el que determina el comportamiento del sistema.

Si graficamos estos polos tendremos:

```python
rs, ks = root_locus(G1)
```

Esta gráfica se le conoce como **Lugar geométrico de las raíces**, y en ella las cruces representan a los polos que obtuvimos, las lineas que parten de estas cruces representan el movimiento de estos polos bajo relalimentación, lo cual vermos en la próxima práctica.

En esta gráfica podemos notar que las raices son complejas, y que su parte real es negativa; esta última característica es la que nos inidca que el comportamiento de este sistema será estable; para corroborar esto podemos simular y gráficar su comportamiento:

```python
ts = linspace(0, 15, 500)
t, y = step_response(G1, ts)

plot(t, y);
```

Por otro lado, si creamos una función de transferencia que tenga polos con parte real $0$, nos dará un comportamiento **Criticamente estable**.

```python
G2 = tf([1], [1, 0, 1])
G2.pole()
```

```python
rs, ks = root_locus(G2)
```

```python
ts = linspace(0, 15, 500)
t, y = step_response(G2, ts)

plot(t, y);
```

O una función de transferencia con polos con parte real positiva, su comportamiento será inestable:

```python
G3 = tf([1], [1, -1, 1])
G3.pole()
```

```python
rs, ks = root_locus(G3)
```

```python
ts = linspace(0, 15, 500)
t, y = step_response(G3, ts)

plot(t, y);
```

Despues de esto, podemos preguntarnos si en lugar de esperar que el comportamiento de nuestro sistema sea el adecuado, podemos construir una función con el comportamiento deseado, y la respuesta es si; podemos por ejemplo utilizar dos raices copmletamente reales para obtener un comportamiento subamortiguado:

$$
G_4 = \frac{1}{s+1} \cdot \frac{1}{s+3} = \frac{1}{s^2 + 4s + 3}
$$

```python
G4 = tf([1], [1, 1])*tf([1], [1, 3])
G4.pole()
```

```python
rs, ks = root_locus(G4)
```

```python
ts = linspace(0, 15, 500)
t, y = step_response(G4, ts)

plot(t, y);
```

O incluso asegurarnos que el comportamiento sea inestable:

```python
G5 = tf([1], [1, -1])*tf([1], [1, 3])
G5.pole()
```

```python
rs, ks = root_locus(G5)
```

```python
ts = linspace(0, 15, 500)
t, y = step_response(G5, ts)

plot(t, y);
```

---

## Ejercicio

* Define una función de transferencia ```G3```, con un comportamiento inestable.

<!-- grade_id: cell-d11c012745f12b1c, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

```python

```

<!-- grade_id: cell-baee0712ec0b9edc, points: 1, grade: true, solution: false -->
```python
assert not all([polo.real<0 for polo in G3.pole()])
```

* Define una función de transferencia ```G4```, con un comportamiento estable.

<!-- grade_id: cell-63152233393e9d37, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-ec93e6a974fb720a, points: 1, grade: true, solution: false -->
```python
assert all([polo.real<0 for polo in G4.pole()])
```

* Define una función de transferencia ```G5```, con un comportamiento criticamente estable.

<!-- grade_id: cell-e78c7ecb7c8ae5ad, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-69573e1b1f783a25, points: 1, grade: true, solution: false -->
```python
assert any([polo.real==0 for polo in G5.pole()])
```

* Define una función de transferencia ```G6```, con un comportamiento sobreamortiguado.

<!-- grade_id: cell-8f42ab3016ff800b, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-bf08e813b5682c0f, points: 1, grade: true, solution: false -->
```python
from numpy import pi, angle
assert all([pi - pi/4 < angle(polo) < pi + pi/4 for polo in G6.pole()])
```

* Define una función de transferencia ```G7```, con un comportamiento subamortiguado.
> Sugerencia: Utiliza como base la función de transferencia del sistema masa - resorte - amortiguador.

<!-- grade_id: cell-8bf539faab8b860f, points: 0, grade: false, solution: true -->
```python
# ESCRIBE TU CODIGO AQUI
raise NotImplementedError
```

<!-- grade_id: cell-944dfbadaaec118a, points: 1, grade: true, solution: false -->
```python
assert not all([pi - pi/4 < angle(polo) < pi + pi/4 for polo in G7.pole()])
```

---